---
title: Plane Fit Problem
exports:
  - format: pdf
    template: plain_latex
    output: exports/02-plane-fit.pdf
    logo: false
    link: true
downloads:
  - file: exports/02-plane-fit.pdf
  - file: 02-PlaneFitProblem.ipynb
math:
    '\calA': '{\cal A}'
    '\calB': '{\cal B}'
    '\calC': '{\cal C}'
    '\calD': '{\cal D}'
    '\calE': '{\cal E}'
    '\calF': '{\cal F}'
    '\calG': '{\cal G}'
    '\calH': '{\cal H}'
    '\calI': '{\cal I}'
    '\calJ': '{\cal J}'
    '\calK': '{\cal K}'
    '\calL': '{\cal L}'
    '\calM': '{\cal M}'
    '\calN': '{\cal N}'
    '\calO': '{\cal O}'
    '\calP': '{\cal P}'
    '\calQ': '{\cal Q}'
    '\calR': '{\cal R}'
    '\calS': '{\cal S}'
    '\calT': '{\cal T}'
    '\calU': '{\cal U}'
    '\calV': '{\cal V}'
    '\calW': '{\cal W}'
    '\calX': '{\cal X}'
    '\calY': '{\cal Y}'
    '\calZ': '{\cal Z}'
    '\setA': '\textsf{A}'
    '\setB': '\textsf{B}'
    '\setC': '\textsf{C}'
    '\setD': '\textsf{D}'
    '\setE': '\textsf{E}'
    '\setF': '\textsf{F}'
    '\setG': '\textsf{G}'
    '\setH': '\textsf{H}'
    '\setI': '\textsf{I}'
    '\setJ': '\textsf{J}'
    '\setK': '\textsf{K}'
    '\setL': '\textsf{L}'
    '\setM': '\textsf{M}'
    '\setN': '\textsf{N}'
    '\setO': '\textsf{O}'
    '\setP': '\textsf{P}'
    '\setQ': '\textsf{Q}'
    '\setR': '\textsf{R}'
    '\setS': '\textsf{S}'
    '\setT': '\textsf{T}'
    '\setU': '\textsf{U}'
    '\setV': '\textsf{V}'
    '\setW': '\textsf{W}'
    '\setX': '\textsf{X}'
    '\setY': '\textsf{Y}'
    '\setZ': '\textsf{Z}'
    '\bfa': '\mathbf{a}'
    '\bfb': '\mathbf{b}'
    '\bfc': '\mathbf{c}'
    '\bfd': '\mathbf{d}'
    '\bfe': '\mathbf{e}'
    '\bff': '\mathbf{f}'
    '\bfg': '\mathbf{g}'
    '\bfh': '\mathbf{h}'
    '\bfi': '\mathbf{i}'
    '\bfj': '\mathbf{j}'
    '\bfk': '\mathbf{k}'
    '\bfl': '\mathbf{l}'
    '\bfm': '\mathbf{m}'
    '\bfn': '\mathbf{n}'
    '\bfo': '\mathbf{o}'
    '\bfp': '\mathbf{p}'
    '\bfq': '\mathbf{q}'
    '\bfr': '\mathbf{r}'
    '\bfs': '\mathbf{s}'
    '\bft': '\mathbf{t}'
    '\bfu': '\mathbf{u}'
    '\bfv': '\mathbf{v}'
    '\bfw': '\mathbf{w}'
    '\bfx': '\mathbf{x}'
    '\bfy': '\mathbf{y}'
    '\bfz': '\mathbf{z}'
    '\bfalpha': '\boldsymbol{\alpha}'
    '\bfbeta': '\boldsymbol{\beta}'
    '\bfgamma': '\boldsymbol{\gamma}'
    '\bfdelta': '\boldsymbol{\delta}'
    '\bfepsilon': '\boldsymbol{\epsilon}'
    '\bfzeta': '\boldsymbol{\zeta}'
    '\bfeta': '\boldsymbol{\eta}'
    '\bftheta': '\boldsymbol{\theta}'
    '\bfiota': '\boldsymbol{\iota}'
    '\bfkappa': '\boldsymbol{\kappa}'
    '\bflambda': '\boldsymbol{\lambda}'
    '\bfmu': '\boldsymbol{\mu}'
    '\bfnu': '\boldsymbol{\nu}'
    '\bfomicron': '\boldsymbol{\omicron}'
    '\bfpi': '\boldsymbol{\pi}'
    '\bfrho': '\boldsymbol{\rho}'
    '\bfsigma': '\boldsymbol{\sigma}'
    '\bftau': '\boldsymbol{\tau}'
    '\bfupsilon': '\boldsymbol{\upsilon}'
    '\bfphi': '\boldsymbol{\phi}'
    '\bfchi': '\boldsymbol{\chi}'
    '\bfpsi': '\boldsymbol{\psi}'
    '\bfomega': '\boldsymbol{\omega}'
    '\bfxi': '\boldsymbol{\xi}'
    '\bfell': '\boldsymbol{\ell}'
    '\bfA': '\mathbf{A}'
    '\bfB': '\mathbf{B}'
    '\bfC': '\mathbf{C}'
    '\bfD': '\mathbf{D}'
    '\bfE': '\mathbf{E}'
    '\bfF': '\mathbf{F}'
    '\bfG': '\mathbf{G}'
    '\bfH': '\mathbf{H}'
    '\bfI': '\mathbf{I}'
    '\bfJ': '\mathbf{J}'
    '\bfK': '\mathbf{K}'
    '\bfL': '\mathbf{L}'
    '\bfM': '\mathbf{M}'
    '\bfN': '\mathbf{N}'
    '\bfO': '\mathbf{O}'
    '\bfP': '\mathbf{P}'
    '\bfQ': '\mathbf{Q}'
    '\bfR': '\mathbf{R}'
    '\bfS': '\mathbf{S}'
    '\bfT': '\mathbf{T}'
    '\bfU': '\mathbf{U}'
    '\bfV': '\mathbf{V}'
    '\bfW': '\mathbf{W}'
    '\bfX': '\mathbf{X}'
    '\bfY': '\mathbf{Y}'
    '\bfZ': '\mathbf{Z}'
    '\bfGamma': '\boldsymbol{\Gamma}'
    '\bfDelta': '\boldsymbol{\Delta}'
    '\bfTheta': '\boldsymbol{\Theta}'
    '\bfLambda': '\boldsymbol{\Lambda}'
    '\bfPi': '\boldsymbol{\Pi}'
    '\bfSigma': '\boldsymbol{\Sigma}'
    '\bfUpsilon': '\boldsymbol{\Upsilon}'
    '\bfPhi': '\boldsymbol{\Phi}'
    '\bfPsi': '\boldsymbol{\Psi}'
    '\bfOmega': '\boldsymbol{\Omega}'
    '\bbA': '\mathbb{A}'
    '\bbB': '\mathbb{B}'
    '\bbC': '\mathbb{C}'
    '\bbD': '\mathbb{D}'
    '\bbE': '\mathbb{E}'
    '\bbF': '\mathbb{F}'
    '\bbG': '\mathbb{G}'
    '\bbH': '\mathbb{H}'
    '\bbI': '\mathbb{I}'
    '\bbJ': '\mathbb{J}'
    '\bbK': '\mathbb{K}'
    '\bbL': '\mathbb{L}'
    '\bbM': '\mathbb{M}'
    '\bbN': '\mathbb{N}'
    '\bbO': '\mathbb{O}'
    '\bbP': '\mathbb{P}'
    '\bbQ': '\mathbb{Q}'
    '\bbR': '\mathbb{R}'
    '\bbS': '\mathbb{S}'
    '\bbT': '\mathbb{T}'
    '\bbU': '\mathbb{U}'
    '\bbV': '\mathbb{V}'
    '\bbW': '\mathbb{W}'
    '\bbX': '\mathbb{X}'
    '\bbY': '\mathbb{Y}'
    '\bbZ': '\mathbb{Z}'
    '\p'  : '\partial'
    '\bfzero': '\mathbb{0}'
---

```python
# Initialize Otter
import otter
grader = otter.Notebook("02-PlaneFitProblem.ipynb")
```

<!-- BEGIN QUESTION -->

# Homework (PlaneFit): Problem 1
(10 marks)

Find the minimum of the following function
$$f(\bfx) = \bfx^\top A^\top A \bfx + 2 \bfb^\top \bfx + c$$

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Homework (Plane fit): Problem 2
(10 marks)

Find the minimum of the following function (assume $\Sigma_1 \in \bbR^{n \times n}$ and $\Sigma_2 \in \bbR^{n \times n}$ to be symmetric and positive semi-definite matrices).

$$f(\bfx) = (\bfx - \bfmu_1)^\top \Sigma_1 (\bfx - \bfmu_1)
+ (\bfx - \bfmu_2)^\top \Sigma_2 (\bfx - \bfmu_2) $$

Assume $\bfx \in \bbR^n$, $\bfmu_1 \in \bbR^n$ and $\bfmu_2 \in \bbR^n$.

$\Sigma$ is not the summation symbol here, it is just another matrix. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Homework (Plane fit): Problem 3
(30 marks)

Least square fit in 2D fits a line to a given set of 2D points. This homework asks you to repeat the minimization procedure to find a plane that fits a 2D points. Below you are provided with helper code to load a point cloud, and visualize the points and the plane. Your task is to find the equation of plane that fits the point cloud. 

1. Show the math and derivation for your code
2. Implement the code in python and numpy to find the coefficients of the equation of plane and write your answer in the space provided below. Submit this ipynb file to brightspace. (30 marks)

```python
%pip install plotly requests
def wget(url, filename):
    """
    Download files using requests package
    """
    import requests
    r = requests.get(url)
    with open(filename, 'wb') as fd:
        for chunk in r.iter_content():
            fd.write(chunk)
wget('https://github.com/wecacuee/ECE490-S24-Neural-Networks/raw/master/notebooks/02-linear-models/table_top.npy',
     'table_top.npy')
```

```python
import plotly.graph_objects as go
from plotly.subplots import make_subplots

def points_on_plane(abc):
    a, b, c = abc
    x = np.linspace(-1, 1, 100, dtype=np.float32)
    y = np.linspace(0, 1, 100, dtype=np.float32)

    x, y = np.meshgrid(x, y)
    z = a*x + b*y + c
    return x, y, z

def visualize_plane(pts, abc, fig=None):
    _fig = fig
    fig = make_subplots() if fig is None else fig
    x, y, z = points_on_plane(abc)
    plt_surface = go.Surface(x=x, y=y, z=z)
    
    fig.add_trace(plt_surface)
    visualize_points(pts, fig=fig)
    if _fig is None:
        fig.show()

def visualize_points(pts, fig=None):
    _fig = fig
    fig = make_subplots() if fig is None else fig
    scatter_trace = go.Scatter3d(x=pts[:, 0], y=pts[:, 1], z=pts[:, 2], mode='markers', marker=dict(size=1))
    fig.add_trace(scatter_trace)
    if _fig is None:  
      fig.show()
```

```python
import numpy as np
pts = np.load('table_top.npy').astype('float32') # Nx3
visualize_points(pts)
pts.min(axis=0), pts.max(axis=0)
```

<!-- grade_id: cell-bb7ef5a91325977a, points: 0, grade: false, solution: true -->
```python
def plane_fit(pts):
    ...
    return abc
```

```python
# Find abc such that z = a x + b y + c
abc = plane_fit(pts)
visualize_plane(pts, abc)
def test_plane_fit(plan_fit, env):
    pts = env['pts']
    abc = plane_fit(pts)
    assert np.abs(pts[:,2] - np.hstack((pts[:, :2], np.ones((pts.shape[0],1)))) @ abc).mean() < 0.4
test_plane_fit(plane_fit, env=globals())
```

```python
grader.check("q8")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q8 (points: 30)
```
>>> abc = plane_fit(pts)
>>> visualize_plane(pts, abc)
```
```
>>> def test_plane_fit(plan_fit, env):
...     pts = env['pts']
...     abc = plane_fit(pts)
...     assert np.abs(pts[:, 2] - np.hstack((pts[:, :2], np.ones((pts.shape[0], 1)))) @ abc).mean() < 0.4
>>> test_plane_fit(plane_fit, env=globals())
```
